<a href="https://colab.research.google.com/github/Kaymuni/class-assignments-/blob/main/01_Exercises.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

## Exercise 01.1

Degrees Fahrenheit ($T_F$) are converted to degrees Celsius ($T_c$) using the formula:

$$
T_c = 5(T_f - 32)/9
$$

Write a program to convert 46 degrees Fahrenheit to degrees Celsius and print the result.
Write your program such that you can easily change the temperature in Fahrenheit that you are converting from.

Use the variable name `T_c` for the temperature in degrees Celsius.

<!-- grade_id: cell-ca2d6af7fabf2918, points: 0, grade: false, solution: true -->
```python
T_f = 46
...
```

<!-- grade_id: cell-a7542313c8cdf4ed, points: 0, grade: true, solution: false -->
```python
## test ##
import math
T_c = 5 * (T_f - 32) / 9
assert math.isclose(T_c, 7.77777777777777)
```

## Exercise 01.2

You have been tasked with developing a mortgage affordability test tool. A component is to compute the monthly
interest payments for a range of scenarios.

Interest on a particular mortgage is charged at fixed rate above the Bank of England (BoE) 'official Bank Rate'.
Interest is computed per annum, and interest payments spread equally over each month of the year.
Write a program that computes the interest payable each month, with variables for:

1. Loan principal (amount borrowed, fixed)
1. Official Bank Rate (percentage, expressed per annum)
1. Rate over the official Bank Rate (percentage, expressed per annum)

Test your program with a loan principal of £172,000, BoE rate of 2.25%, and rate over the
BoE rate of 1.49%. Use the variable name `interest` for the monthly interest amount.

<!-- grade_id: cell-aaab53918bb9627d, points: 0, grade: false, solution: true -->
```python
# Assign data to variables
principal = 172000

...
```

Output:
```
Ellipsis
```

<!-- grade_id: cell-258e3a2c80f6a765, points: 0, grade: true, solution: false -->
```python
## test ##
import math

# Define the rates as per the exercise description
boe_rate = 2.25 # Official Bank Rate (percentage)
rate_over_boe = 1.49 # Rate over BoE (percentage)

# Calculate total annual interest rate in decimal
total_annual_rate_decimal = (boe_rate + rate_over_boe) / 100

# Calculate annual interest amount
annual_interest = principal * total_annual_rate_decimal

# Calculate monthly interest amount
interest = annual_interest / 12

assert math.isclose(interest, 536.0666666666667)
```

```python
import sympy as sp

def solve_second_order_ode():
    # Define symbols
    x = sp.Symbol('x')
    y = sp.Function('y')

    print("--- Second Order ODE Solver (ay'' + by' + cy = f(x)) ---")

    try:
        # Input coefficients and the function f(x)
        a = float(input("Enter coefficient a: "))
        b = float(input("Enter coefficient b: "))
        c = float(input("Enter coefficient c: "))
        f_input = input("Enter f(x) (e.g., 'exp(x)', 'sin(x)', or '0'): ")
        f = sp.sympify(f_input)

        # Define the ODE: a*y''(x) + b*y'(x) + c*y(x) = f(x)
        ode = sp.Eq(a * y(x).diff(x, x) + b * y(x).diff(x) + c * y(x), f)

        # 1. Find the Homogeneous Solution (y_h)
        # Solve the ODE with f(x) = 0
        ode_h = sp.Eq(a * y(x).diff(x, x) + b * y(x).diff(x) + c * y(x), 0)
        sol_h = sp.dsolve(ode_h, y(x))
        y_h = sol_h.rhs

        # 2. Find the Particular Solution (y_p)
        # Full solution minus the homogeneous part
        full_sol_gen = sp.dsolve(ode, y(x))
        # To isolate y_p, we set the arbitrary constants C1 and C2 to 0
        y_p = full_sol_gen.rhs.subs({sp.Symbol('C1'): 0, sp.Symbol('C2'): 0})

        print(f"\nHomogeneous Solution (y_h): {y_h}")
        print(f"Particular Solution (y_p): {y_p}")
        print(f"General Solution: y(x) = {full_sol_gen.rhs}")

        # 3. Handle Auxiliary Conditions
        print("\n--- Input Auxiliary Conditions ---")
        x0 = float(input("Enter x value for conditions (x0): "))
        y0 = float(input(f"Enter y({x0}): "))
        dy0 = float(input(f"Enter y'({x0}): "))

        # Solve for constants C1 and C2
        # y(x0) = y0
        cond1 = sp.Eq(full_sol_gen.rhs.subs(x, x0), y0)
        # y'(x0) = dy0
        cond2 = sp.Eq(full_sol_gen.rhs.diff(x).subs(x, x0), dy0)

        constants = sp.solve((cond1, cond2), (sp.Symbol('C1'), sp.Symbol('C2')))

        # Final Solution
        final_sol = full_sol_gen.rhs.subs(constants)

        print("\n--- Results ---")
        print(f"Constants found: {constants}")
        print(f"Final Solution: y(x) = {final_sol}")

        return y_h, y_p, final_sol

    except Exception as e:
        print(f"An error occurred: {e}")

# Run the function
if __name__ == "__main__":
    solve_second_order_ode()
```

Output:
```
--- Second Order ODE Solver (ay'' + by' + cy = f(x)) ---
Enter coefficient a: 1
Enter coefficient b: -5
Enter coefficient c: 6
Enter f(x) (e.g., 'exp(x)', 'sin(x)', or '0'): 12*exp(x)

Homogeneous Solution (y_h): C1*exp(2.0*x) + C2*exp(3.0*x)
Particular Solution (y_p): 6.0*exp(x)
General Solution: y(x) = C1*exp(2.0*x) + C2*exp(3.0*x) + 6.0*exp(x)

--- Input Auxiliary Conditions ---
Enter x value for conditions (x0): 0
Enter y(0.0): 10
Enter y'(0.0): 0

--- Results ---
Constants found: {C1: 18.0000000000000, C2: -14.0000000000000}
Final Solution: y(x) = 6.0*exp(x) + 18.0*exp(2.0*x) - 14.0*exp(3.0*x)
```

```python

```